---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

Apply a random regressor. You can get this from scikit lego. 

`@instructions`
Apply a random regressor. You can get this from scikit lego.

`@hint`
Apply a random regressor. You can get this from scikit lego.

`@pre_exercise_code`
```{python}
from sklearn.dummy import RandomRegressor
```

`@sample_code`
```{python}

```

`@solution`
```{python}
RandomRegressor()
```

`@sct`
```{python}

```
